---
title: Test
description: Test
---

## Adding two numbers

```yaml
type: NormalExercise
key: ee931e74d8
xp: 100
```

In this section you will add two numbers in Python called a and b

`@instructions`
Add two numbers

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
print(___+___)
```

`@solution`
```{python}
print(3+4)
```

`@sct`
```{python}
7
```
