---
title: test
description: test2
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 00f45e8c9e
xp: 100
```

Now we're multiplying two numbers

`@instructions`
Multiply a and b whereby a=4 and b=3

`@hint`


`@pre_exercise_code`

```{python}
a=1
b=2
```


`@sample_code`

```{python}
print(a*b)
```


`@solution`

```{python}
print(a*b)
```


`@sct`

```{python}
12
```
